# Introduction to Machine Learning: Supervised Learning

**Instructor:** Daniel Acuna, Ph.D.
**Position:** Associate Professor of Computer Science
**Institution:** University of Colorado Boulder

---

Lab 2: Regression model evaluation

---

## Setup (do not edit)

```python
import pathlib
from typing import Tuple, Dict, List

import numpy as np
import pandas as pd
from sklearn.model_selection import train_test_split
from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_squared_error, r2_score

RANDOM_STATE: int = 42  # global seed for full reproducibility
np.random.seed(RANDOM_STATE)

_DATA_PATH = pathlib.Path("auto.csv")
if not _DATA_PATH.exists():
    raise FileNotFoundError(
        "auto.csv is missing from the lab directory. Please download it or ask the TA "
        "for assistance."
    )
```

## 1. Load the Dataset *(5 points)*

Write a function `load_auto_data()` that reads `auto.csv` into a `pandas.DataFrame`. Store the shape of the DataFrame in a variable called **`q1_shape`**.

<!-- grade_id: load-auto, points: 0, grade: false, solution: true -->
```python
def load_auto_data() -> pd.DataFrame:
    """Load the Auto dataset from ``auto.csv``.

    Returns
    -------
    pd.DataFrame
        A DataFrame containing the auto dataset.
    """
    load_auto_data = pd.read_csv(_DATA_PATH)
    return load_auto_data




# Compute the answer required by the autograder
q1_shape: Tuple[int, int] = load_auto_data().shape
```

<!-- grade_id: load-auto-test, points: 5, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 5 points
# Test Cell: Question 1
assert isinstance(q1_shape, tuple), "q1_shape must be a tuple"
assert len(q1_shape) == 2, "q1_shape should have 2 elements (rows, cols)"
assert q1_shape[0] > 0 and q1_shape[1] > 0, "Shape values must be positive"
print(f"Dataset shape: {q1_shape}")
```

Output:
```
Dataset shape: (392, 8)
```

## 2. Prepare Data *(10 points)*

Create the feature matrix `X` and target vector `y`.
- `X` should contain the columns: `displacement`, `horsepower`, `weight`, `acceleration`.
- `y` should contain the `mpg` column.
Store the number of features in **`q2_num_features`**.

<!-- grade_id: prepare-data, points: 0, grade: false, solution: true -->
```python
df = load_auto_data()
X = df.drop(columns=["cylinders", "mpg", "year", "origin"])
y = df["mpg"]
q2_num_features: int = X.shape[1]

# your code here
# raise NotImplementedError
```

<!-- grade_id: prepare-data-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 2
assert "X" in locals() and "y" in locals(), "X and y must be defined."
assert isinstance(X, pd.DataFrame), "X should be a DataFrame"
assert isinstance(y, pd.Series), "y should be a Series"
assert isinstance(q2_num_features, int), "q2_num_features should be an integer"
assert q2_num_features > 0, "Number of features must be positive"
print(f"Number of features: {q2_num_features}")
```

Output:
```
Number of features: 4
```

## 3. Train/Test Split *(10 points)*

Perform an 80/20 split of `X` and `y`. Use `random_state=RANDOM_STATE`. Store the row counts of each split in a tuple **`q3_split_counts = (n_train, n_test)`**.

<!-- grade_id: split-data, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
random_state = RANDOM_STATE
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=random_state)
n_train = X_train.shape[0]
n_test = X_test.shape[0]
q3_split_counts = (n_train, n_test)
# your code here
# raise NotImplementedError
```

<!-- grade_id: split-data-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 3
assert isinstance(q3_split_counts, tuple), "q3_split_counts must be a tuple"
assert len(q3_split_counts) == 2, "q3_split_counts should have 2 elements"
train_rows, test_rows = q3_split_counts
assert train_rows > 0 and test_rows > 0, "Both splits must have positive counts"
assert train_rows > test_rows, "Training set should be larger than test set"
print(f"Train rows: {train_rows}, Test rows: {test_rows}")
```

Output:
```
Train rows: 313, Test rows: 79
```

## 4. Train Linear Regression Model *(15 points)*

Train a `LinearRegression` model on the training data. Store the fitted model in a variable named **`model`**.

<!-- grade_id: train-model, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
model = LinearRegression()
model.fit(X_train, y_train)
# your code here
# raise NotImplementedError
```

Output:
```
LinearRegression()
```

<!-- grade_id: train-model-test, points: 15, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 15 points
# Test Cell: Question 4
assert "model" in locals(), "The 'model' variable is not defined."
assert hasattr(model, "fit"), "The model should have a 'fit' method."
assert hasattr(model, "predict"), "The model should have a 'predict' method."
assert hasattr(model, "coef_"), "The model has not been fitted yet."
print(f"Model trained successfully with {len(model.coef_)} coefficients")
```

Output:
```
Model trained successfully with 4 coefficients
```

## 5. Calculate Test MSE *(15 points)*

Use the trained model to make predictions on the test set, then compute the Mean Squared Error (MSE). Store the result in **`q5_mse`** (float, rounded to three decimals).

<!-- grade_id: calc-mse, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
q5_mse: float = float(round(mean_squared_error(y_test, model.predict(X_test)), 3))
```

<!-- grade_id: calc-mse-test, points: 15, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 15 points
# Test Cell: Question 5
assert isinstance(q5_mse, float), "MSE must be a float."
assert q5_mse >= 0, "MSE must be non-negative."
assert q5_mse < 1000, "MSE seems unreasonably large. Check your calculation."
print(f"Test MSE: {q5_mse:.3f}")
```

Output:
```
Test MSE: 18.066
```

## 6. Calculate Test R² *(15 points)*

Compute the R-squared (R²) value for the test set. Store the result in **`q6_r2`** (float, rounded to three decimals).

<!-- grade_id: calc-r2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
q6_r2: float = float(round(r2_score(y_test, model.predict(X_test)), 3))
```

<!-- grade_id: calc-r2-test, points: 15, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 15 points
# Test Cell: Question 6
assert isinstance(q6_r2, float), "R-squared must be a float."
assert q6_r2 <= 1.0, "R-squared cannot exceed 1.0."
assert q6_r2 > -10, "R-squared value seems unreasonable. Check your calculation."
print(f"Test R²: {q6_r2:.3f}")
```

Output:
```
Test R²: 0.646
```

## 7. Reusable Evaluation Function *(20 points)*

Implement a function **`evaluate_model(model, X_test, y_test)`** that returns a tuple containing the (MSE, R²) score, with each value rounded to three decimals.

<!-- grade_id: eval-func, points: 0, grade: false, solution: true -->
```python
def evaluate_model(model, X_test, y_test):
    y_pred = model.predict(X_test)
    mse = float(round(mean_squared_error(y_test, y_pred), 3))
    r2 = float(round(r2_score(y_test, y_pred), 3))
    return mse, r2
```

<!-- grade_id: eval-func-test, points: 20, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 20 points
# Test Cell: Question 7
test_mse, test_r2 = evaluate_model(model, X_test, y_test)
assert isinstance(test_mse, float) and isinstance(
    test_r2, float
), "Function must return two floats."
assert test_mse >= 0, "MSE must be non-negative."
assert test_r2 <= 1.0, "R-squared cannot exceed 1.0."
print(f"Function returned MSE: {test_mse:.3f}, R²: {test_r2:.3f}")
```

Output:
```
Function returned MSE: 18.066, R²: 0.646
```

## 8. Analyze Coefficients *(10 points)*

Identify the feature with the largest *negative* coefficient (i.e., the one that most negatively impacts MPG). Store its name as a string in **`q8_strongest_negative_feature`**.

<!-- grade_id: analyze-coef, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
q8_strongest_negative_feature: str = X.columns[np.argmin(model.coef_)]
# your code here
#raise NotImplementedError
```

<!-- grade_id: analyze-coef-test, points: 10, grade: true, solution: false -->
```python
# If all tests pass (there might be hidden tests), you will earn 10 points
# Test Cell: Question 8
assert isinstance(q8_strongest_negative_feature, str), "Answer must be a string."
assert len(q8_strongest_negative_feature) > 0, "Feature name cannot be empty."
assert (
    q8_strongest_negative_feature in X.columns
), "Feature name must be one of the columns in X."
print(f"Feature with most negative coefficient: {q8_strongest_negative_feature}")
```

Output:
```
Feature with most negative coefficient: acceleration
```

## Next Steps

Congratulations on completing the assignment! Before submitting:

1. Make sure all your cells run without errors.
2. Ensure you've answered all parts of each question.
3. If any autograder tests fail, revisit your answers.
